![DSA log](dsalogo-Abuja.png)

### Instructions

1. Make sure you are using a version of notebook greater than v.3. If you installed Anaconda with python 3 - this is likely to be fine. The next piece of code will check if you have the right version.
2. The notebook has both some open test cases that you can use to test the functionality of your code - however it will be run on another set of test cases that you can't from which marks will be awarded. So passing all the tests in this notebook is not a guarantee that you have done things correctly - though its highly probable.
3. Also make sure you submit a notebook that doesn't return any errors. One way to ensure this is to run all the cells before you submit the notebook.
4. When you are done create a zip file of your notebook and upload that
5. For each cell where you see "YOUR CODE HERE" delete the return notImplemented statement when you write your code there - don't leave it in the notebook.
6. Once you are done, you are done.

## Python Basics

<!-- grade_id: cell-0fc4baabfcf40810, points: 1, grade: true, solution: false -->
```python
# Import some stuff that you will use in this notebook
from nose.tools import assert_equal
import numpy as np
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

#### Question 1

Complete the function to output a string which is the concatenation of the inputs to the function

<!-- grade_id: cell-cd8089abf6c58ed7, points: 0, grade: false, solution: true -->
```python
def get_string(x,y):
    # YOUR CODE HERE
    c = str(x) + str(y)
    return c
```

<!-- grade_id: cell-056d72fe6275f378, points: 1, grade: true, solution: false -->
```python
assert_equal(get_string(4,5), '45')
assert_equal(get_string('hello','world'), 'helloworld')
```

#### Question 2

Complete the function to output the value of $y$ given $x$ and $p$. Function definition is $y = 4 + 5x^p$ <br>
**Condition:** If the $x$ is not positive the value of y should be 0

<!-- grade_id: cell-3468a2b044f2bc45, points: 0, grade: false, solution: true -->
```python
def ypolynomial(x,p):
    # YOUR CODE HERE
    if x<0:
        y=0
    else:
        y=4+(5*(x**p))
    return(y)
    
```

<!-- grade_id: cell-f557c11ac2b4dcf9, points: 2, grade: true, solution: false -->
```python
assert_equal(ypolynomial(3,3), 139)
assert_equal(ypolynomial(-2,3), 0)
```

#### Question 3

Complete the function to returns <i>'obese'</i> or <i>'not obese'</i> given weight (in Kg) and and height (in meters).  Note: one is obesse if the Body Mass Index (BMI) is greater or equal to 30 and $BMI = \frac{weight}{height^2}$

<!-- grade_id: cell-0b75e3e42e7dc6ea, points: 0, grade: false, solution: true -->
```python
def isObesse(weight, height):
    # YOUR CODE HERE
    bmi=weight/(height**2)
    if bmi>=30:
        return("obese")
    else:
        return("not obese")
    
```

<!-- grade_id: cell-fc0d92d6e95d3056, points: 2, grade: true, solution: false -->
```python
assert_equal(isObesse(80,1.58), 'obese')
assert_equal(isObesse(122,2.03), 'not obese' )

```

#### Question 4

Complete the function to return a list of numbers between 0 and a number $N$ that are perfectly divisible by $p$ (with out leaving a remainder). <br>
**Hint**: $N$ should also be inclusive in the numbers being considered. If N is negative use N == 20

<!-- grade_id: cell-f60c36c52f5521ff, points: 0, grade: false, solution: true -->
```python
def divisibles(N, p):
    # YOUR CODE HERE
    perfect=[]
    if N>=0:
        for i in range(0,N+1):
            if i % p==0:
                perfect.append(i)
    elif N<0:
        N=20
        for i in range(0,N+1):
            if i% p==0:
                perfect.append(i)
    return perfect
       
```

<!-- grade_id: cell-3b0136dc3012d384, points: 2, grade: true, solution: false -->
```python
assert_equal(divisibles(10,2), [0,2,4,6,8,10])
assert_equal(divisibles(-10,2), [0,2,4,6,8,10,12,14,16,18,20])
```

#### Question 5
Complete the function to output the list, mylist, with duplicate values removed. 

<!-- grade_id: cell-c53babe05da1c87e, points: 0, grade: false, solution: true -->
```python
def removeDuplicates(mylist):
    # YOUR CODE HERE
    return list(set(mylist))
```

<!-- grade_id: cell-04d54a4e147911ae, points: 2, grade: true, solution: false -->
```python
assert_equal(removeDuplicates([1,3,2,3,4]), [1,2,3,4])
assert_equal(removeDuplicates([4,1,4,3]), [1,3,4])
```

#### Question 6

Complete the function to output the sum of the last and first numbers in the list. Note if list is empty return 0.

<!-- grade_id: cell-767f49260faf257d, points: 0, grade: false, solution: true -->
```python
def sumOfFirstAndLast(listx ):
    # YOUR CODE HERE
    sum=[]
    if len(list(listx)) == 0:
        return 0
    else:
        for i in listx:
            sum=listx[0]+listx[-1]
        return sum
```

<!-- grade_id: cell-7f6c59d3fbfd68a1, points: 2, grade: true, solution: false -->
```python
assert_equal(sumOfFirstAndLast([3,8,2,12]), 15)
assert_equal(sumOfFirstAndLast([4]), 8)

```

#### Below is an sample code that reverses digits in a number. Modify this code, if you find it useful, to to attempt questions 7, 8, 9 and 10. 

```python
def reverseNumber(number):

    revnumber = 0
    while(number>0):
        lastdigit = number%10
        revnumber = (revnumber*10)+lastdigit;
        number = number//10
        
    return revnumber
```

#### Question 7

Complete the function to output the sum of digits of the input number.

<!-- grade_id: cell-5151049547ad3052, points: 0, grade: false, solution: true -->
```python
def sumDigits(number):
    # YOUR CODE HERE
    sumnum = 0
    while number>0:
        lastdigit= number % 10
        sumnum = sumnum+lastdigit
        number= number// 10
    return sumnum
```

<!-- grade_id: cell-b8753a3bc70a9af4, points: 2, grade: true, solution: false -->
```python
assert_equal(sumDigits(1234),10)
assert_equal(sumDigits(172034),17)
```

#### Question 8

Complete the function to output the sum of even number digits of the input number.

<!-- grade_id: cell-bbe62c9bb21ac35d, points: 0, grade: false, solution: true -->
```python
def countEvenDigits(number):
    
    # YOUR CODE HERE
    sumnum=[]
    evenDigit=[]
    while number>0:
        digit= number % 10
        sumnum = digit
        if sumnum % 2 ==0:
            evenDigit.append(sumnum)
        number= number// 10
    return len(evenDigit)
```

<!-- grade_id: cell-dd57add680b55a32, points: 2, grade: true, solution: false -->
```python
assert_equal(countEvenDigits(189156782),4)
assert_equal(countEvenDigits(431728564),5)
```

#### Question 9

Complete the function to output a dictionary with both the sum of even and odd digits with the
keys 'even' and 'odd' respectively

<!-- grade_id: cell-695a2d2fdc6a4b67, points: 0, grade: false, solution: true -->
```python
def countOddEvenDigits(number):
    
    # YOUR CODE HERE
    oddDigit=[]
    evenDigit=[]
    
    while number>0:
        digit= number % 10
        sumnum = digit
        if sumnum % 2 ==0:
            evenDigit.append(sumnum)
        else:
            oddDigit.append(sumnum)            
        number= number// 10
    numDict={'odd':len(oddDigit),'even':len(evenDigit)}
    return numDict

        
```

<!-- grade_id: cell-82c5f1e7e19418d6, points: 2, grade: true, solution: false -->
```python
assert_equal(countOddEvenDigits(189156782),{'odd':5,'even':4})
assert_equal(countOddEvenDigits(431728564),{'odd':4,'even':5})
```

#### Question 10


The concept of <a href="https://en.wikipedia.org/wiki/Check_digit">check digit</a> is widely used in identification codes including bank account numbers, International Standard Book Numbers (ISBNs) and Universal Product Codes (UPCs). Check digits are useful in reducing manual input errors in digital applications, as they the codes or numbers with check digits can be validated without any other extra information. This question will focus on ISBN 13 that uses the European Article Number (EAN) 13 barcode specification, described in the image below 

![ISBN 13 Calculation  Example](ISBN13example.png) [Extracted from Wikipedia]

Using the information above complete the function to determine the validity of the ISBN13 and returns 'Valid' or 'Invalid for a valid or invalid ISBN13 respectively.


<!-- grade_id: cell-12046f0d6b5bdb38, points: 0, grade: false, solution: true -->
```python
def verify_ISBN13(ISBN):
    ISBNList=list(str(ISBN))
    ISBNarrayOrig=np.array([list(str(ISBN))])
    ISBNarray=np.array([list(str(ISBN))], dtype='int')
    ISBNarrayPos= np.delete(ISBNarray, -1)
    
    
    Odd=(ISBNarrayPos[::2])
    Even=(ISBNarrayPos[1::2])
    
    
    SumOdd=np.sum(Odd)
    SumEven=np.sum(Even)
    
    NewEven= SumEven*3
    partial_sum = SumOdd  + NewEven

    result = 10 -(partial_sum % 10)

    if (int(result) == int(ISBNList[-1])):
        return ('Valid')
    else:
        return ('Invalid')
   
```

<!-- grade_id: cell-47c64fad5b6e1655, points: 2, grade: true, solution: false -->
```python
assert_equal(verify_ISBN13(4006381333931), 'Valid')
assert_equal(verify_ISBN13(9788175257665), 'Valid')
assert_equal(verify_ISBN13(9788175257645), 'Invalid')
```